```python
# Initialize Otter
import otter
grader = otter.Notebook("06_datentypen.ipynb")
```

# Computational Thinking - WS 2023/2024

## Datentypen

Autor: Prof. Dr.-Ing. Benedikt Dietrich

### Aufgabe 1

Schreiben Sie Code, der zwei Variablen `x` und `y` anlegt. Weisen Sie `x` den Wert 5 zu und `y` den Wert 2.5. Addieren Sie die beiden Variablen und speichern Sie den Wert in einer neuen Variable namens `z`. 

Lassen Sie sich den Datentyp von allen drei Variablen in folgendem Format ausgeben:

```text
  Datentyp von x: ...
  Datentyp von y: ...
  Datentyp von z: ...
```

Was können Sie beobachten?

```python
# Put your code here

...
```

```python
grader.check("q1")
```

### Aufgabe 2

Prüfen Sie, was passiert, wenn Sie zu einer Variable vom Datentyp `string` eine Variable vom Typ `int` addieren.

```python
...
```

```python
x = 5
my_text = "Der Wert von x ist: "
# BEGIN SOLUTION
result = my_text + str(x)
# END SOLUTION
```

## Abgabe

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert x == 5, 'x hat nicht den geforderten Wert'
```
```
>>> assert y == 2.5, 'y hat nicht den geforderten Wert'
```
```
>>> assert z == x + y, 'y hat nicht den geforderten Wert'
```
